Tactile localization of the breast, areola, and nipple

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

The manuscript incorrectly describes the result as poor spatial acuity. Acuity measures the average absolute error, and acuity is good when response biases are absent. Precision relates to the error variance. It is common to see high precision with low acuity or vice versa. Just noticeable differences assess precision or spatial resolution, while points of subjective equality evaluate acuity or bias. Similar confusions between these terms appear throughout the manuscript.

While I do not agree with the reviewer's usage of the word “acuity” and a cursory Google search does not agree with the provided definition, I have replaced acuity with precision as appropriate to improve clarity.

A paragraph within the next section seems to follow up on this insight by examining the across-participant consistency of the differences in tactile spatial resolution between body parts. To this aim, pairwise rank correlations between body sites are conducted. This analysis raises red flags from a statistical point of view. 1) An ANOVA and its follow-up tests assume no variation in the size of the tested effect but varying base values across participants. Thus, if significant differences between conditions are confirmed by the original statistical analysis, most participants will have better spatial resolution in one condition than the other condition, and the difference between body sites will be similar across participants. 2) Correlations are power-hungry, and non-parametric tests are power-hungry. Thus, the number of participants needed for a reliable rank correlation analysis far exceeds that of the study. In sum, a correlation should emerge between body sites associated with significantly different tactile JNDs; however, these correlations might only be significant for body sites with pronounced differences due to the sample size.

We have entirely removed this result from both the text and supplement.

The data do not support this conclusion. The conclusion that the nipple is perceived as a unit is based on poor tactile localization performance for touches on the nipple compared to the areola. The problem is that the localization task is a quadrant identification task with the center being at the nipple. Quadrants for the areola could be significantly larger due to the relative size of the areola and the nipple; the results section seems to suggest this was accounted for when placing the tactile stimuli within the quadrants, but the methods section suggests otherwise. Additionally, the areola has an advantage because of its distance from the nipple, which leads to larger Euclidean distances between the centers of the quadrants than for the nipple. Thus, participants should do better for the areola than for the nipple even if both sites have the same tactile resolution.

We agree with this interpretation and have updated the language throughout.

Categorization accuracy in each area was tested against chance using a Monte Carlo test, which is fine, though the calculation of the test statistic, Z, should be reported in the Methods section, as there are several options. Localization accuracies are then compared between areas using a paired t-test. It is a bit confusing that once a distribution-approximating test is used, and once a test that assumes Gaussian distributions when the data is Bernoulli/Binomial distributed. Sampling-based and t-tests are very robust, so these surprising choices should have hardly any effect on the results.

Excellent point. We have replaced the paired t-test with a signed rank test and added text to the methods to expand upon this.

A correlation based on N=4 participants is dangerously underpowered. A quick simulation shows that correlation coefficients of randomly sampled numbers are uniformly distributed at such a low sample size. This likely spurious correlation is not analyzed, but quite prominently featured in a figure and discussed in the text, which is worrisome.

We have removed this panel to reduce this concern.

The conclusion that tactile percepts are drawn toward the nipple is based on localization biases for tactile stimuli on the breast compared to the back. Unfortunately, the way participants reported the tactile locations introduces a major confound. Participants indicated the perceived locations of the tactile stimulus on 3D models of these body parts. The nipple is a highly distinctive and cognitively represented landmark, far more so than the scapula, making it very likely that responses were biased toward the nipple regardless of the actual percepts. One imperfect but better alternative would have been to ask participants to identify locations on a neutral grey patch and help them relate this patch to their skin by repeatedly tracing its outline on the skin.

While I wholeheartedly agree with the sentiments of the reviewer, in our experience performing these tests across many women we have found that the variability of the morphology of the breast makes it incredibly hard for women to perform this task in the way the reviewer is describing. Consequently, there is likely no perfect version of the task. That said, we have endeavored to acknowledge the limitations of the approach in the discussion.

Participants also saw their localization responses for the previously touched locations. This is unlikely to induce bias towards the nipple, but it renders any estimate of the size and variance of the errors unreliable. Participants will always make sure that the marked locations are sufficiently distant from each other.

I again respectfully disagree with this interpretation. If the participants were to always make sure marked locations were sufficiently distant from each other then the degree of error and bias would be similar between regions given that the visual pattern would be almost identical. As this is not true in the data, I disagree with the premise, though we hope the changes to the discussion acknowledge limitations with the data collection method.

Null-hypothesis significance testing only lets scientists either reject the null hypothesis or not. The latter does NOT mean the Null hypothesis is true, i.e., it can never be concluded that there is no effect. This rule applies to every NHST test. However, it raises particular concerns with distribution tests. The only conclusion possible is that the data are unlikely from a population with the tested distribution; these tests do not provide insight into the actual distribution of the data, regardless of whether the result is significant or not.

Thank you for this comment. We have updated the language to make it explicit that we do not mean to imply failing to deviate from the Null distribution does not mean that they are in fact Null in nature.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

I am wondering whether the interpretation of "the nipple as a sensory unit" is also supported by localization performance as reported in the analysis around Fig. 3 and supplementary Fig. 2. I cannot really see the error lines in that figure, and cannot tell whether any of the touches were on the nipple proper. Specifically I am wondering whether touch to the nipple is reliably attributed to the nipple, and touch to the areola to the areola, or whether confusion exists between the two. The description of the nipple as a sensory unit implies reliable attribution of touch to the respective area. Also the discussion (lines 309ff) is ambiguous about this.

Thank you for this comment. We have removed language about the nipple being a unit and reframed the text in the discussion. We have also clarified that touches were indeed on the nipple.

typos etc.

lines 68-71 - implied causality is not backed up by evidence and could be the other way around than stated here

line 82 grammar is inconsistent

lines 199-200, "on the nipple" occurs twice

Thank you for catching these. We have addressed the typos and grammar. We have also added a citation to the sentence where this exact hypothesis is stated. We have also relaxed the language to imply it is indeed a hypothesis.

Changes from the earlier version to the revised version

Abstract

Earlier version: Touch plays a key role in our perception of our body and shapes our interactions with the world, from the objects we manipulate to the people we touch. While the tactile sensibility of the hand has been extensively characterized, much less is known about touch on other parts of the body. Despite the important role of the breast in lactation as well as in affective and sexual touch, relatively little is known about its sensory properties. To fill this gap, we investigated the spatial acuity of the breast and compared it to that of the hand and back, body regions that span the range of tactile spatial acuity. First, we found that the tactile acuity of the breast was even lower than that of the back, heretofore the paragon of poor acuity. Second, acuity was lower for larger breasts, consistent with the hypothesis that innervation capacity does not scale with body size. Third, touches to different regions of the nipple were largely indistinguishable, suggesting that the nipple is a sensory unit. Fourth, localization errors were systematically biased toward the nipple.

Revised version: Touch plays a key role in our perception of our body and shapes our interactions with the world, from the objects we manipulate to the people we touch. While the tactile sensibility of the hand has been extensively characterized, much less is known about touch on other parts of the body. Despite the important role of the breast in lactation as well as in affective and sexual touch, relatively little is known about its sensory properties. To fill this gap, we investigated the ability of women to locate touches on the breast and compared it to that of the hand and back, body regions that span the range of tactile discriminative capabilities. First, we found that the tactile precision of the breast was even lower than that of the back, heretofore the paragon of poor precision. Second, precision was lower for breasts that had undergone greater expansion, consistent with the hypothesis that innervation capacity does not scale with body size. Third, touches to different regions of the nipple were largely indistinguishable, suggesting sparse innervation density. Fourth, localization errors were systematically biased toward the nipple.

Introduction

Earlier version: The sense of touch fulfills a variety of different functions in everyday life, from guiding our interactions with the environment to supporting affective communication and sexual function (1). One of the key properties of touch sensations is that they are localized to a specific part of the body: contact on the shoulder produces a sensation experienced on the shoulder, for example. The precision with which we can localize events on the skin has been shown to be determined by the innervation density at that skin location (2, 3). Because the skin of the fingertips and lips is the most densely innervated, the acuity of these body regions is highest, conferring to us an enhanced ability to distinguish touches on the fingers and lips even when they are close to each other. Innervation density and function are intrinsically related: the fingertips are densely innervated because they account for the vast majority of contacts with objects, and precise information about object interactions is critical to dexterous manipulation (4). In contrast, the acuity on the back is low because precise localization of a touch on the back is of limited value. Skin surface area also plays a role in innervation density; for example, people with large hands exhibit lower acuity than those with small ones (5, 6). This phenomenon is hypothesized to reflect the fact that the number of nerve fibers does not scale with body size: a large body will be more sparsely innervated than a small one given a fixed number of nerve fibers.

Revised version: The sense of touch fulfills a variety of different functions in everyday life, from guiding our interactions within the environment to supporting affective communication and sexual function (1). One of the key properties of touch sensations is that they are localized to a specific part of the body: contact on the shoulder produces a sensation experienced on the shoulder, for example. The precision with which we can localize events on the skin has been shown to be determined by the innervation density at that location (2, 3). Because the skin of the fingertips and lips is the most densely innervated, the precision of these body regions is highest, conferring to us an enhanced ability to distinguish touches on the fingers even when they are close to each other. Innervation density and function are intrinsically related: the fingertips are thought to be densely innervated because they account for the vast majority of contacts with objects and precise information about object interactions is critical to dexterous manipulation (4, 5). In contrast, the precision on the back is low because precise localization of a touch on the back is of limited value. Skin surface area also plays a role in innervation density; for example, people with large hands exhibit lower precision than those with small ones (6, 7). This phenomenon is hypothesized to reflect the fact that the number of nerve fibers does not scale with body size: a large body will be more sparsely innervated than a small one given a fixed number of nerve fibers.

Earlier version: While tactile acuity has been extensively studied on the limbs and face, acuity on the torso has received far less experimental attention (7, 8), with the breast being largely ignored beyond its anatomy (9, 10). To fill this gap, we sought to characterize the spatial acuity of the female breast which has roles in lactation, affective touch, and sex. These functions set it apart from other regions of the body. In previous studies, the tactile acuity of the breast was found to be comparably low to that of the back and calf (8). However, the experimental approaches have since received scrutiny to their susceptibility to inaccuracies (11); only the outer breast was tested (excluding the nipple-areolar complex, NAC), and the relationship with breast size was not assessed, or only examined the NAC but not the outer breast (9). Given that breasts grow to vastly different sizes relatively late in life after the nervous system is nearly fully developed (12), the female breast offers a powerful test of the fixed innervation hypothesis, which would predict that people with larger breasts would have poorer spatial acuity than people with smaller ones.

Revised version: While tactile precision has been extensively studied on the limbs and face, precision on the torso has received far less experimental attention (8, 9), with the breast being largely ignored beyond its anatomy (10, 11). To fill this gap, we sought to characterize the spatial precision of the female breast which has roles in lactation, affective touch, and sex. These functions set it apart from other regions of the body. In previous studies, the tactile precision of the breast was found to be comparably low to that of the back and calf (9). However, the experimental approaches have received scrutiny to their susceptibility to inaccuracies (12); only the outer breast was tested (excluding the nipple-areolar complex, NAC), and the relationship with breast size was not assessed, or only examined the NAC but not the outer breast (10). Given that the timeline for breast development extends well past that of nervous system development (13), the female breast offers a powerful test of the fixed innervation hypothesis, which would predict that people with larger breasts would have poorer spatial precision than people with smaller ones.

Earlier version: In the present study, we first measured the tactile acuity of two regions of the breast in women– the outer (lateral) breast and medial breast, which includes the NAC and compared these to the hand and back. Second, we examined the relationship between the acuity of the outer breast and breast size. Third, we examined women’s ability to judge the absolute location of touches to their breast. We first found that the spatial acuity of the breast is very low, with similar sensitivity to the back. Second, spatial acuity is inversely correlated with breast size, as predicted from the fixed innervation hypothesis. Third, touches to different parts of the nipple are indistinguishable, suggesting that the nipple is a sensory unit. Fourth, touches on the outer breast are systematically mis-localized as being closer to the nipple than they actually are.

Revised version: In the present study, we first measured the tactile precision of two regions of the breast in women– the outer (lateral) breast and medial breast, which includes the NAC and compared these to the hand and back. Second, we examined the relationship between the precision of the outer breast and breast size. Third, we examined women’s ability to judge the absolute location of touches to their breast. We first found that the spatial precision of the breast is very low, with similar sensitivity to the back. Second, spatial precision is inversely correlated with breast size, as predicted from the fixed innervation hypothesis. Third, touches to different parts of the nipple are indistinguishable, suggesting sparse tactile innervation. Fourth, touches on the outer breast are systematically mis-localized as being closer to the nipple than they actually are.

Results

Earlier version: The breast has low spatial acuity

Revised version: The breast has low spatial precision

Earlier version: We first measured each participant’s ability to judge the relative position of two touches applied in succession at two nearby locations on the skin of the hand, back, and breast using a punctate probe. The first of the two touches (the reference) was at the same position on each trial, and the second touch (the comparison) was either above or below the first at a pre-specified distance. For the hand, distances ranged from 1 to 10 mm; for the other body regions, distances ranged from 2.5 to 40 mm, anticipating lower acuity based on preliminary testing (Figure 1). The participant’s task was to report whether the test stimulus was located above or below the reference stimulus. We then assessed the participant’s performance as a function of the distance between the test and reference and characterized the distance from the reference required to reliably locate the test stimulus – the just noticeable difference (JND) – for each participant and region.

Revised version: We first measured each participant’s ability to judge the relative position of two touches applied in succession at two nearby locations on the skin of the hand, back, and breast using a punctate probe. The first of the two touches (the reference) was at the same position on each trial, and the second touch (the comparison) was either above or below the first at a pre-specified distance. For the hand, distances ranged from 1 to 10 mm; for the other body regions, distances ranged from 2.5 to 40 mm, anticipating lower precision based on preliminary testing (Figure 1). The participant’s task was to report whether the test stimulus was located above or below the reference stimulus. We then assessed the participant’s performance as a function of the distance between the test and reference and characterized the distance from the reference required to reliably locate the test stimulus – the just noticeable difference (JND) – for each participant and region.

Location discrimination.

Earlier version: A|Example psychometric functions for one subject on the location discrimination task for the hand, back, outer breast, and medial breast (including NAC). Negative values denote test points “below” the reference and the JND indicates the distance at which the subject could reliably (75% trials) locate the stimulus. B| Distribution of JNDs for each subject at each region. C| Relationship between the difference in size between the bust and underbust (Δ Bust) and the JND of the lateral breast. D| Relationship between the JND for the hand and breast across participants.

Revised version: A| Example psychometric functions for one subject on the location discrimination task for the hand, back, outer breast, and medial breast (including NAC). Negative values denote test points “below” the reference and the JND indicates the distance at which the subject could reliably (75% trials) locate the stimulus. B| Distribution of JNDs for each subject at each region. C| Relationship between the difference in size between the bust and underbust (Δ Bust) and the JND of the lateral breast and D| medial breast, respectively.

Earlier version: As expected, the hand demonstrated significantly greater spatial acuity (median [25th,75th] percentiles: 2.16 [1.73, 2.93] mm) than either the back (5.72 [4.36, 8.16] mm) or either breast region (6.63 [5.41, 7.73] mm and 8.00 [5.72, 10.80] mm, for the lateral and medial breast respectively) as indicated by the lower JNDs (Figure 1A, B, 1-way ANOVA: F[3,126] = 22.7, p < 0.01; Tukey’s HSD post-hoc t-test: p < 0.01). Interestingly, the medial breast yielded even shallower functions and lower spatial acuity than did the back (Figure 1A, B, 1-way ANOVA with Tukey’s HSD post-hoc t-test: p = 0.0284), though no statistical differences were observed between the lateral breast and either the back or medial breast (p = 0.3974 and 0.7377, respectively). In other words, touches needed to be between 3 and 4 times as far apart on the breast than on the hand to yield equivalent location discrimination performance.

Revised version: As expected, the hand demonstrated significantly greater spatial precision (median [25th,75th] percentiles: 2.16 [1.73, 2.93] mm) than either the back (5.72 [4.36, 8.16] mm) or either breast region (6.63 [5.41, 7.73] mm and 8.00 [5.72, 10.80] mm, for the lateral and medial breast respectively) as indicated by the lower JNDs (Figure 1A, B, 1-way ANOVA: F[3,126] = 22.7, p < 0.01; Tukey’s HSD post-hoc t-test: p < 0.01). Interestingly, the medial breast yielded even shallower functions and lower spatial precision than did the back (Figure 1A, B, 1-way ANOVA with Tukey’s HSD post-hoc t-test: p = 0.0284), though no statistical differences were observed between the lateral breast and either the back or medial breast (p = 0.3974 and 0.7377, respectively). In other words, touches needed to be between 3 and 4 times as far apart on the breast than on the hand to yield equivalent location discrimination performance.

Earlier version: Tactile acuity is worse for women with larger breasts

Revised version: Tactile precision is worse for women with larger breasts

Earlier version: The spatial acuity of the hand has been shown to depend on the size of the hand, with smaller hands yielding better acuity (5). With this observation in mind, we investigated whether the inter-participant differences in breast acuity might be driven in part by differences in breast size. To determine the relative expansion (increase in surface area) of each participant’s breast, we computed the difference between their bust (circumference of the torso at the nipple line) and underbust (circumference of the torso at the inframammary fold). Comparing this value with their spatial acuity, we found that the acuity of the lateral breast decreased (JNDs increased) as breast size increased (Figure 1C, Pearson’s correlation with Bonferroni PHC (n=4): r = 0.734, p < 0.01), but found no effect at any other region (Supplementary Figure 1A-C, p > 0.05 for all), implying that any expansion is limited to the skin of the lateral breast and that the observed relationship between breast size and acuity was not spurious.

Revised version: The spatial precision of the hand has been shown to depend on the size of the hand, with smaller hands yielding better precision (6). With this observation in mind, we investigated whether the inter-participant differences in breast precision might be driven in part by differences in breast size. To determine the relative expansion (increase in surface area) of each participant’s breast, we computed the difference between their bust (circumference of the torso at the nipple line) and underbust (circumference of the torso at the inframammary fold). Comparing this value with their spatial precision, we found that the precision of the lateral breast decreased (JNDs increased) as breast size increased (Figure 1C, Pearson’s correlation with Bonferroni PHC (n=4): r = 0.734, p < 0.01), a phenomenon that did not extend to the medial breast (Figure 1C), hand, or back (Supplementary Figure 1, p > 0.05 for all), implying that any expansion is preferentially limited to the skin of the lateral breast and that the observed relationship between breast size and precision was not spurious.

Added in the revised version: Tactile events are poorly discriminated on the nipple

Removed in the revised version: To examine if participants demonstrated consistent relative acuity across regions, we first ranked the acuity performance of each participant within region and then took the average rank across regions to determine the overall acuity rank of each participant. Measuring the rank-correlation between the overall rank and that of each region, we found that acuity of the hand, lateral, and medial breast were all significantly correlated with overall rank (Kendall’s correlation with Bonferroni PHC: p < 0.01 for all) but that back acuity was not (p = 0.13 after PHC). While this observation implies that the relative skill or attention of each participant influenced the acuity of each region, differences in the spatial expansion of a region could likely override this effect. Thus, we also examined if the acuity of each region was predictive of any other by comparing their JNDs. Consistent with the fixed innervation hypothesis, we found no relationship between the spatial acuity of each region, however, we found that the JNDs of the hand and medial breast were strongly related (Figure 1D, Pearson’s correlation, r = 0.685, p < 0.01; r = −0.007 to 0.375, p = 0.25 – 0.99 for other comparisons, Supplementary Figure 1D-H; Bonferroni PHC used for all where n = 6), though no efforts were made to measure hand size to examine if the two were related. Consequently, the spatial expansion that occurs at the medial breast is likely similar to that of the hand, implying that expansion of the NAC may be independent of the remainder of the breast.

Removed in the revised version: The nipple is a unit

Earlier version: While the result that the medial breast has the lowest spatial acuity is not necessarily surprising given the lack of mechanoreceptors near the surface of the skin (13), the magnitude of the correlation was unexpected. Indeed, we observed that the JNDs for each participant were, on average, only marginally smaller than the measured diameter of the nipple (80% ± 52%). This would imply that 2 points on opposite halves of the nipple (80% diameter apart) might be confused as the same point. Given that the NAC is inherently a heterogeneous structure as it is composed of three anatomically distinct regions of the breast (the nipple, the areola, and the nearby outer breast), we reasoned that this approach may be biased and thus not reflect the functional ability to locate sensations on the breast.

Revised version: While the result that the medial breast has the lowest spatial resolution is not necessarily surprising given the lack of mechanoreceptors near the surface of the skin (14), the magnitude of the correlation was unexpected. Indeed, we observed that the JNDs for each participant were, on average, only marginally smaller than the measured diameter of the nipple (80% ± 52%). This would imply that 2 points on opposite halves of the nipple (80% diameter apart) might be confused as the same point. Given that the NAC is inherently a heterogeneous structure as it is composed of three anatomically distinct regions of the breast (the nipple, the areola, and the nearby outer breast), we reasoned that this approach may be biased and thus not reflect the functional ability to locate sensations on the breast.

Earlier version: To assess if participants could reliably perceive tactile events across the nipple-areola complex, we delivered punctate touches to each quadrant on the nipple or areola and the participant reported the quadrant in which the touch had been delivered (Figure 2A). In particular, we marked the edge of the nipple or the 50% point between the center of the nipple and the edge of the areola in each of the cardinal directions such that the proportional distance was constant across participants. When touching different quadrants of the areola, participants were able to reliably locate stimuli (Figure 2B, mean ± standard deviation: 55% ± 21% where chance is 25%, one-sided Monte-Carlo test, Z = 15, p < 0.01), with 8 of the 10 participants performing significantly above chance (p < 0.05 after Bonferroni PHC where n = 10). Of the remaining participants, one performed marginally above chance (40%, p = 0.25 after Bonferroni PHC) while the other systematically misperceived the location and performed well below chance (7.5%, p = 1 after Bonferroni PHC). On the nipple, however, participants were consistently worse at locating stimuli on the nipple than the breast (paired t-test, t = 3.42, p < 0.01) where only 3 of the 10 participants outperformed chance, though the group as a whole outperformed chance (Figure 2B, 36% ± 13%; Z = 5.5, p < 0.01). This finding suggests that the ability to locate tactile events is significantly worse at the nipple than at the areola and re-affirms our earlier finding that spatial acuity at the medial breast is poor such that contact events to it can only be poorly localized at best. Finally, to determine if this task was also influenced by spatial expansion, we sought to compare the breast size with performance on the quadrant task. Although we found negative trends for performance at both the areola and nipple, the data was collected post-hoc and only 4 of the 10 participants returned for measurements, thereby precluding statistical analysis (Figure 2C). This negative trend implicitly contradicts the previous result where one might expect equal performance regardless of size as the location of the stimuli was scaled to the size of the nipple and areola. However, given the absence of a reference point, systematic biases are more likely to occur and thus this may reflect a relationship between localization bias and breast size.

Revised version: To assess if participants could reliably perceive tactile events at the level of either the areola or nipple, we delivered punctate touches to each quadrant on the nipple or areola and the participant reported the quadrant in which the touch had been delivered (Figure 2A). Specifically, we marked the edge of the nipple or the 50% point between the center of the nipple and the edge of the areola in each of the cardinal directions such that the proportional distance was constant across participants. When touching different quadrants of the areola, participants were able to reliably locate stimuli (Figure 2B, mean ± standard deviation: 55% ± 21% where chance is 25%, one-sided Monte-Carlo test, Z = 15, p < 0.01), with 8 of the 10 participants performing significantly above chance (p < 0.05 after Bonferroni PHC where n = 10). Of the remaining participants, one performed marginally above chance (40%, p = 0.25 after Bonferroni PHC) while the other systematically misperceived the location and performed well below chance (7.5%, p = 1 after Bonferroni PHC). On the nipple, however, participants were consistently worse at locating stimuli than the areola (Wilcoxon signed-rank test, p = 0.0137) where only 3 of the 10 participants outperformed chance, though the group as a whole outperformed chance (Figure 2B, 36% ± 13%; Z = 5.5, p < 0.01). This finding suggests that the ability to locate tactile events is significantly worse at the nipple than at the areola, consistent with equivalent innervation and the smaller size of the nipple, and re-affirms our earlier finding that spatial precision at the medial breast is poor. Finally, to determine if this task was also influenced by spatial expansion, we sought to compare the breast size with performance on the quadrant task. Unfortunately, only 4 of the 10 participants returned for measurement and while negative trends were observed, the sample size is insufficient for formal analysis.

Earlier version: Spatial acuity of the nipple and areola.

Revised version: Spatial discrimination at the nipple and areola.

Earlier version: A|Test locations for the quadrant discrimination task. The subject reported location using a number from 1 to 4 progressing clockwise from “above.” B| Proportion correct for participant at each test location. Solid markers indicate participants whose performance was significantly above chance (N=10). C| Relationship between breast size (bust – underbust) and performance on the quadrant task on the areola or nipple.

Revised version: A| Test locations for the quadrant discrimination task. The outer ring represents the areola and the inner ring the nipple. The subject reported location using a number from 1 to 4 progressing clockwise from “above.” B| Proportion correct for participant at each test location. Solid markers indicate participants whose performance was significantly above chance (N=10).

Localization of tactile events on the breast is biased towards the nipple

In the experiments described above, we investigated the participants’ ability to distinguish the relative locations of two touches to the breast. Next, we examined their ability to identify the absolute location of a tactile stimulus on their torso. To this end, we presented a single touch to the breast or the back with a punctate probe at one of 25 locations after which each participant marked the location of the touch on a three-dimensional digital image of the participant’s own breast or back (Figure 3A, B).

Absolute localization of contact events on the breast and back.

A| Example localization task data for the breast or B| back of one participant. Gray surface represents scanned torso. Black crosses indicate the true location of each stimulus, purple lines indicate the vector between the stimulus location and reported location, and blue lines indicate the vector between the stimulus location and average reported location across blocks. C| Reporting error (3D Euclidean distance) as a function of distance across participants for the breast and D| back. Error indicates mean error for each trial (mean of purple vector length in A,B), bias indicates the error of the average response (blue vectors in A,B), and imprecision is the mean pairwise error between individual responses for each point (distance between the end of purple vectors in A,B). E| The distribution of angles (2D, no depth axis) between the stimulus location and the reported location (gray) or the difference between said angle and the angle towards the relevant landmark (red) for the breast (nipple) and F| back (scapula).

Earlier version: To estimate the degree to which participants could accurately localize touch events, we computed the error for each body part across participants and found that, contrary to the acuity task, participants had lower errors on their breasts in comparison to their back (Supplementary Figure 2A; paired t-test, t[9] = 7.4, p < 0.001). Next, we computed this error as a function of the distance from the center landmark for each region (nipple or the inferior angle of the scapula, i.e. the bottom corner, Figure 3C, D) and found that in the case of the breast, the errors increased systematically as the distance from the nipple increased consistent with the quadrant localization task (Linear Mixed Model, LMM: distance: n = 248, B = 0.234, t = 9.194, p < 0.001, Table 1). This association was not found for the back (distance: n = 248, B = 0.093, t = 1.714, p = 0.087, Table 2). Then, using an LMM to compare the groups, we found both a significant interaction between the effect of distance and the location (location x distance: n = 496, B = −0.139, t = −2.723, p = 0.006, Table 3). Given that any errors are likely a result of two sources of noise: bias and response variability (imprecision), we next sought to quantify the relative contribution of the two sources. We estimated the bias for each point by computing the centroid of all reports for a given stimulus across trials and measured the error between the centroid and the stimulus location. To measure the imprecision, we computed the mean pairwise distance between each of the reported locations for a given stimulus location and the mean of those reported locations (Figure 3C, D). We found that, in both the breast and back, the error that originated from systematic bias significantly outweighed that of the imprecision (location x error type: n = 995, B = −23.05, t = 12.728, p < 0.001).

Revised version: To estimate the degree to which participants could accurately localize touch events, we computed the error for each body part across participants and found that, contrary to the precision task, participants had lower errors on their breasts in comparison to their back (Supplementary Figure 2A; paired t-test, t[9] = 7.4, p < 0.001). Next, we computed this error as a function of the distance from the center landmark for each region (nipple or the inferior angle of the scapula, i.e. the bottom corner, Figure 3C, D) and found that in the case of the breast, the errors increased systematically as the distance from the nipple increased consistent with the quadrant localization task (Linear Mixed Model, LMM: distance: n = 248, B = 0.234, t = 9.194, p < 0.001, Table 1). This association was not found for the back (distance: n = 248, B = 0.093, t = 1.714, p = 0.087, Table 2). Then, using an LMM to compare the groups, we found both a significant interaction between the effect of distance and the location (location x distance: n = 496, B = -0.139, t = -2.723, p = 0.006, Table 3). Given that any errors are likely a result of two sources of noise: bias and response variability (imprecision), we next sought to quantify the relative contribution of the two sources. We estimated the bias for each point by computing the centroid of all reports for a given stimulus across trials and measured the error between the centroid and the stimulus location. To measure the imprecision, we computed the mean pairwise distance between each of the reported locations for a given stimulus location and the mean of those reported locations (Figure 3C, D). We found that, in both the breast and back, the error that originated from systematic bias significantly outweighed that of the imprecision (location x error type: n = 995, B = -23.05, t = 12.728, p < 0.001).

Earlier version: As bias was prevalent in the localization task, we next sought to determine if the bias was systematic across stimulus locations. Consequently, we assessed the tendency for the error vector to be uniformly distributed in two dimensions (Supplementary Figure 2B-C). Importantly, because the back is a relatively flat surface in comparison to the breast, error vectors were only computed in the horizontal plane and depth was excluded when computing the angular error. When computing the absolute angle between stimulus location and the average reported location (blue vectors in Figure 3A, B) and computed the distribution across participants (Figure 3E, F). Across participants, we found no consistent excursions in the absolute biases from a normal distribution for the breast or back (permutation test for uniformity: R = 0.1314, p = 0.9331 and R = 0.1856, p = 0.7171, respectively). Next, when we assessed if the bias with respect to the center point (nipple or scapula) was consistent, we found that both the breast and back tended to exhibit biases (R = 0.3140, p = 0.0100 and R = 0.6164, p < 0.0001, respectively), with the effect roughly doubled for the breast.

Revised version: As bias was prevalent in the localization task, we next sought to determine if the bias was systematic across stimulus locations. Consequently, we assessed the tendency for the error vector to be uniformly distributed in two dimensions (Supplementary Figure 2B-C). Importantly, because the back is a relatively flat surface in comparison to the breast, error vectors were only computed in the horizontal plane and depth was excluded when computing the angular error. When computing the absolute angle between stimulus location and the average reported location (blue vectors in Figure 3A, B) and computed the distribution across participants (Figure 3E, F). Across participants, the distribution of biases did not significantly deviate from uniformity (permutation test for uniformity: R = 0.1314, p = 0.9331 and R = 0.1856, p = 0.7171, respectively). Next, when we assessed if the bias with respect to the center point (nipple or scapula) was consistent, we found that both the breast and back tended to exhibit biases (R = 0.3140, p = 0.0100 and R = 0.6164, p < 0.0001, respectively), with the effect roughly doubled for the breast.

Earlier version: To further quantify this observation, we computed the vector strength – a measure of circular uniformity - of the biases for individual participants. For the breast we found that the absolute angle of the biases tended to be uniformly distributed across participants, though some participants demonstrated biases (e.g. towards the top of the shoulder, Supplementary Figure 2D-F; vector strength Monte-Carlo test: p < 0.05 for 4/10 participants). Examining the back, there was a significant tendency for all but 1 participant to be biased towards a given direction, however, the direction was inconsistent across participants implying that reports were not biased towards any single landmark. Next, we computed the distribution of relative angles between the bias vector and the center point of our task and found that 7 of 10 participants were biased towards the nipple (vector strength Monte-Carlo: p < 0.05) and 2 were biased towards the scapula, consistent with the notion that biases are significant but inconsistent.

Revised version: To further quantify this observation, we computed the vector strength – a measure of circular uniformity-of the biases for individual participants. For the breast we did not detect significant non-uniform distributions across participants, though some participants demonstrated biases (e.g. towards the top of the shoulder, Supplementary Figure 2D-F; vector strength Monte-Carlo test: p < 0.05 for 4/10 participants). Examining the back, significant non-uniform distributions were observed for all but 1 participant, however, the direction was inconsistent across participants implying that reports were not biased towards any single landmark. Next, we computed the distribution of relative angles between the bias vector and the center point of our task and found that 7 of 10 participants were biased towards the nipple (vector strength Monte-Carlo: p < 0.05) and 2 were biased towards the scapula, suggesting that while individual biases are present, they vary substantially across the population.

Earlier version: Given these findings, we conclude that the breast has lower tactile acuity than the hand and is instead comparable to the back. Moreover, localization of tactile events to both the back and breast are inaccurate but localizations to the breast are consistently biased towards the nipple.

Revised version: Given these findings, we conclude that the breast has lower tactile precision than the hand and is instead comparable to the back. Moreover, localization of tactile events to both the back and breast are inaccurate but localizations to the breast are consistently biased towards the nipple.

Discussion

Earlier version: First, we found that the spatial acuity of the lateral breast – excluding the nipple and areola – is almost four times lower than that of the hand, and even lower than that of the back, previously considered the epitome of poor acuity. Second, the nipple is essentially a sensory unit, as touches to different aspects of the nipple are nearly indistinguishable from one another. Third, the acuity of the breast tends to be poorer for women with large breasts, consistent with the theory that innervation capacity is fixed. Together, the data indicates that the nipple constitutes a landmark on the breast, as evidenced by the fact that absolute localization judgments are less accurate and more biased for touches that are far from the nipple and perceived location is pulled systematically toward the nipple.

Revised version: First, we found that the spatial precision of the lateral breast – excluding the nipple and areola – is almost four times lower than that of the hand, and even lower than that of the back, previously considered the epitome of poor precision. Second, the nipple has such low tactile precision that touches to different aspects of the nipple are nearly indistinguishable from one another. Third, the precision of the breast tends to be poorer for women with large breasts, consistent with the theory that innervation capacity is fixed. Together, the data indicates that the nipple constitutes a landmark on the breast, as evidenced by the fact that absolute localization judgments are less accurate and more biased for touches that are far from the nipple and the perceived location is pulled systematically toward the nipple.

Earlier version: The poor spatial acuity of the breast

Revised version: The poor spatial precision of the breast

Earlier version: The poor spatial acuity of the breast – about four times lower than that of the hand and slightly worse than the back – replicates previous findings (8) achieved using less reliable methods (two point threshold and a ‘same-different’ paradigm). Tactile spatial acuity is determined by density of innervation: more neural tissue is devoted to more highly innervated body regions and this increased central representation is a key contributor to the increased acuity. The low spatial acuity of the nipple and areola is consistent with a histological study revealing these regions to be sparsely innervated (13).

Revised version: The poor spatial precision of the breast – about four times lower than that of the hand and slightly worse than the back – replicates previous findings (9) achieved using less reliable methods (two point threshold and a ‘same-different’ paradigm). Tactile spatial precision is determined by density of innervation: more neural tissue is devoted to more highly innervated body regions, and this increased central representation is a key contributor to the increased precision. The low spatial precision of the nipple and areola is consistent with a histological study revealing these regions to be sparsely innervated (14).

Earlier version: Larger breasts confer lower acuity

Revised version: Larger breasts confer lower precision

Earlier version: We found a significant relationship between breast size and spatial acuity: women with larger breasts tended to exhibit lower spatial acuity on their breasts, consistent with previous findings that tactile acuity scales with body size. Indeed, the spatial acuity of the hand has been shown to depend on the size of the hand, with smaller hands exhibiting better acuity (5, 6). These results are consistent with the hypothesis that the number of tactile nerve fibers does not scale with body size, so fibers are more sparsely distributed on bigger bodies, leading to lower acuity. Consequently, the acuity of the breast is likely determined initially by torso acuity and then any expansion.

Revised version: We found a significant relationship between breast size and spatial precision: women with larger breasts tended to exhibit lower spatial precision on their breasts, consistent with previous findings that tactile precision scales with body size. Indeed, the spatial precision of the hand has been shown to depend on the size of the hand, with smaller hands exhibiting better precision (6, 7). These results are consistent with the hypothesis that the number of tactile nerve fibers does not scale with body size, so fibers are more sparsely distributed on bigger bodies, leading to lower precision. Consequently, the precision of the breast is likely determined initially by torso precision and then as a result of subsequent expansion.

Reports are biased towards the nipple

Earlier version: The mental representation of the body is not uniform and veridical (14). Localization tends to be more precise when stimuli are applied near anatomical points of reference that form perceptual anchor points (15). For example, the navel and spine act as anchor points along the abdomen (16, 17): touches to the navel or spine are never mistaken for touches anywhere else on the abdomen. Furthermore, touches to locations near these two anchor points are mislocalized following a bias toward these areas: touches near the navel are pulled toward the navel and touches near the spine are pulled toward the spine. Similar biases are observed on the back of the hand, where touches are mislocalized to be closer to the fingers than they actually are (18), and on the forearm, where they are pulled toward the wrist (19). Analogously, we found that acuity is highest near the nipple, and perceived locations are pulled toward the nipple, suggesting that the nipple plays a pivotal role in the mental representation of the breast.

Revised version: The mental representation of the body is not uniform and veridical (15). Localization tends to be more precise when stimuli are applied near anatomical points of reference that form perceptual anchor points (16). For example, the navel and spine act as anchor points along the abdomen (17, 18): touches to the navel or spine are never mistaken for touches anywhere else on the abdomen. Furthermore, touches to locations near these two anchor points are mislocalized following a bias toward these areas: touches near the navel are pulled toward the navel and touches near the spine are pulled toward the spine. Similar biases are observed on the back of the hand, where touches are mislocalized to be closer to the fingers than they actually are (19), and on the forearm, where they are pulled toward the wrist (20). Analogously, we found that precision is highest near the nipple, and perceived locations are pulled toward the nipple, suggesting that the nipple plays a pivotal role in the mental representation of the breast. It must be noted, however, that these same biases may be inadvertently influenced by our study design. As participants were annotating 3D meshes of their breasts, the cognitive importance of the nipple may have also caused them to report sensations as closer to the nipple regardless of the actual percepts. This observation motivates standard reporting methodologies for morphologically diverse body parts such as the(21).

Implications for breast prostheses

Earlier version: Understanding the spatial acuity of the breast is particularly important given a recent proliferation of efforts - including using autologous tissue, synthetic grafts and neuroprosthetic or bionic approaches – to restore sensation to the breast following mastectomy (20, 21). In particular, this work shows that the spatial resolution of an implantable sensor sheet should depend on the distance from the nipple. Individual sensors should be placed in each quadrant of the nipple-areolar complex while subsequent sensors should be placed radially with increasing separation (0.5 to 2 cm inter-sensor spacing) up to a 5 cm radius at which point resolution would remain constant at approximately 2 cm. These values, however, depend on the spatial resolution of the stimulation technology which can vary significantly for both peripheral nerve stimulation (22) and intracortical microstimulation (23).

Revised version: Understanding the spatial precision of the breast is particularly important given a recent proliferation of efforts - including using autologous tissue, synthetic grafts and neuroprosthetic or bionic approaches - to restore sensation to the breast following mastectomy (22, 23). In particular, this work shows that the spatial resolution of an implantable sensor sheet should depend on the distance from the nipple. Individual sensors should be placed in each quadrant of the nipple-areolar complex while subsequent sensors should be placed radially with increasing separation (0.5 to 2 cm inter-sensor spacing) up to a 5 cm radius at which point resolution would remain constant at approximately 2 cm. These values, however, depend on the spatial resolution of the stimulation technology which can vary significantly for both peripheral nerve stimulation (24) and intracortical microstimulation (25). Moreover, little is understood about the tactile coding of the primary afferents of the breast and further research will be needed to both inform stimulus patterning and inform safe stimulation parameters (26–28).

Conclusion

Earlier version: The breast has unique sensory properties because (1) it mediates lactation and nursing, (2) it comprises distinct regions – the nipple, areola, and outer breast – which differ in the type of skin and patterns of sensory innervation, (3) it has erectile function and gives rise to erogenous sensations, and (4) it undergoes variable expansion across individuals during puberty. First, we find that spatial acuity on the breast is lower than that of the back, previously regarded as the body region with lowest tactile acuity. Second, spatial acuity is lower in larger breasts, presumably due to sensory innervation being fixed prior to expansion during puberty. Third, the nipple itself is perceived as a single spatial unit and serves as an anchor for the mental map of the breast.

Revised version: The breast has unique sensory properties because (1) it mediates lactation and nursing, (2) it comprises distinct regions – the nipple, areola, and outer breast – which differ in the type of skin and patterns of sensory innervation, (3) it has erectile function and gives rise to erogenous sensations, and (4) it undergoes variable expansion across individuals during puberty. First, we find that spatial precision on the breast is lower than that of the back, previously regarded as the body region with lowest tactile precision. Second, spatial precision is lower in larger breasts, presumably due to sensory innervation being fixed prior to expansion during puberty. Third, the nipple itself has poor precision and yet plays a major role in how participants perceive tactile events on their breast.

Methods

Participants

A total of 48 healthy adult women (24.2 ± 3.1, 19-32 years) participated in this study: Thirty-four in the location discrimination task (mean ± standard deviation, age range; 22.6 ± 2.7, 19-28 years), ten in the medial breast quadrant localization task (25.3 ± 3.8, 19-32 years), and ten in the absolute localization task (24.7 ± 2.9, 20-28 years). Some participants performed multiple tasks. Experimental procedures were performed in accordance with the relevant guidelines and regulations and were approved by the Institutional Review Board of the University of Chicago (IRB 18-0135). The informed consent process included documentation of written consent from each participant, and the participants were compensated for their participation. We excluded any women who, by self-report, were currently pregnant or breastfeeding, had a history of breast surgery, or any diagnosis of neurological illness.

Breast measurements

In addition to self-reported bra size, we collected standardized, objective measurements from each participant. Specifically, we measured the bust – i.e. the circumference of the chest at the level of the nipple – and subtracted from it the under-bust, the circumference at the level of the inframammary fold (mean ± standard deviation, range; 12.4 ± 4.6, 7.6-25.4 cm). We also measured the diameter of the areola and nipple (areola: 37.4 ± 11.9, 21-80 mm; nipple: 12.7 ± 2.7, 7-20 mm).

Tactile Stimuli

Earlier version: Touches were delivered manually with a 1-mm tipped XP-Pen (XP-Pen USA, CA, USA) which allowed the experimenter to lightly press the stimulus against the skin and ensure that the skin was indented uniformly across touches. The XP-Pen was selected instead of monofilaments because, in pilot experiments, several participants reported discomfort with the sharp edges of the monofilaments on their breasts while also allowing for the application of a constant force. The participants were asked to report any discomfort with the application of the stimulus or the inability to feel the application of the stimulus reliably. Note that spatial acuity is consistent across stimulus amplitudes as long as the touch is sufficiently above threshold (24), which was the case here.

Revised version: Touches were delivered manually with a 1-mm tipped XP-Pen (XP-Pen USA, CA, USA) which allowed the experimenter to lightly press the stimulus against the skin and ensure that the skin was indented uniformly across touches. The XP-Pen was selected instead of monofilaments because, in pilot experiments, several participants reported discomfort with the sharp edges of the monofilaments on their breasts while also allowing for the application of a constant force. The participants were asked to report any discomfort with the application of the stimulus or the inability to feel the application of the stimulus reliably. Note that spatial precision is consistent across stimulus amplitudes as long as the touch is sufficiently above threshold (29), which was the case here.

Psychophysical tasks

Location discrimination

Earlier version: Acuity was tested at each of four body regions: the lateral breast, the medial breast (which includes the NAC), the thenar eminence of the hand, and the upper back. The participant – wearing a gown that exposed only the location to be tested – lay supine on a massage table for testing on the lateral breast, medial breast, and thenar eminence, and prone for the testing on the back. Unless the participant expressed a preference, the side of the body to be tested (left/right) was chosen randomly (through a random number generator) but each participant was tested on the same side for all regions. The order in which the different regions were tested was randomized to eliminate any effects of fatigue or learning.

Revised version: Precision was tested at each of four body regions: the lateral breast, the medial breast (which includes the NAC), the thenar eminence of the hand, and the upper back. The participant – wearing a gown that exposed only the location to be tested – lay supine on a massage table for testing on the lateral breast, medial breast, and thenar eminence, and prone for the testing on the back. Unless the participant expressed a preference, the side of the body to be tested (left/right) was chosen randomly (through a random number generator) but each participant was tested on the same side for all regions. The order in which the different regions were tested was randomized to eliminate any effects of fatigue or learning.

On each trial, two touches were applied to nearby locations and the participant’s task was to indicate whether the second touch was above or below the first by pressing one of two buttons on a keypad. The location of one of the two touches (the reference) was consistent across each experimental block and the location of the second (the comparison) varied from trial to trial (each at a pre-specified distance). Fifteen comparison locations, aligned along an axis parallel to the body’s axis, were tested on the back, lateral breast, and medial breast. Touch locations were drawn on the body to ensure repeatable presentation and comparisons were placed at locations 0, 2.5, 5, 7.5, 10, 15, 20 or 40 mm above or below the reference (Figure 1A). Each comparison was repeated five times per block over three blocks (for a total of 15 repeats per comparison, with 225 total trials per body region) in a randomized order. Thirteen comparison locations were tested on the hand, 0, 1, 3, 6, 7.5, and 10 mm away from the reference in both directions along an axis in line with the thumb For the hand, a judgment of “above” indicated that the comparison was displaced toward the thumb relative to the reference. Each location was repeated 5 times across three experimental blocks (for a total of 15 repeats per comparison, with 195 total trials).

Performance was gauged by the proportion of times the participant judged a comparison as above the reference as a function of distance from the reference, where negative distances indicate comparison locations below the reference. Psychometric functions were fit to the data and used to compute the just noticeable difference.

Quadrant discrimination

The objective of this experiment was to assess the degree to which women can distinguish touches to different parts of their NAC. On each trial, the participant was touched at one of four locations on the nipple or areola (organized in a quadrant) and verbally identified the touch location using a number from 1 to 4. Each location was touched ten times per block (40 trials per block, with 80 total trials per breast region). For the areola, each touch was located halfway between the edge of the areola and the nipple. For the nipple, each touch was delivered at the edge of the nipple, in line with each point on the areola (Figure 2A). The testing side was chosen randomly for each participant (6 left/4 right).

Absolute localization

Earlier version: The objective of this experiment was to gauge the accuracy with which women could report where on their breast or back a touch was delivered. On the breast, touches were arranged such that the nipple was the central point and other touches radiated outwards from it (Figure 3A). On the back, touches were arranged the same way around a central point located on the shoulder blade (Figure 3B). First, the participant’s skin was marked, then a three-dimensional scan of the breast or back was obtained using the EM3D application (Brawny Lads Software, LLC.) and .obj files were uploaded to Blender. The surface of the breast was then masked with a uniform layer of grey to obscure the markings. A laptop (13” screen) was positioned such that the participant could use the trackpad to interact with the 3D image (zoom, rotate, drag object) to report where the touch was experienced on the 3D rendering of their breast. Notably, in both cases landmarks were observable on the 3D models (e.g. nipple, shoulder, scapula, vertebral line). Each of the 25 stimulus locations was presented in a pseudorandomized order while the participants closed their eyes. After the touch, the participant opened their eyes and marked the perceived location of the touch on the 3D rendering of their breast or back. The participant was encouraged to manipulate the 3D model freely to obtain the best view of their breast or back on each trial. The participant moved the cursor and clicked to indicate where they perceived the touch. Reported locations accumulated on the 3D model throughout the block (during which each location was touched once.) to encourage placement of markers in accurate relative positions but were removed at the start of each of three blocks (yielding a total of 3 repeats for each of the 25 locations).

Revised version: The objective of this experiment was to gauge the accuracy with which women could report where on their breast or back a touch was delivered. On the breast, touches were arranged such that the nipple was the central point and other touches radiated outwards from it (Figure 3A). Importantly, the two center-most touches occurred on the nipple itself. On the back, touches were arranged the same way around a central point located on the shoulder blade (Figure 3B). First, the participant’s skin was marked, then a three-dimensional scan of the breast or back was obtained using the EM3D application (Brawny Lads Software, LLC.) and .obj files were uploaded to Blender. The surface of the breast was then masked with a uniform layer of grey to obscure the markings. A laptop (13” screen) was positioned such that the participant could use the trackpad to interact with the 3D image (zoom, rotate, drag object) to report where the touch was experienced on the 3D rendering of their breast. Notably, in both cases landmarks were observable on the 3D models (e.g. nipple, shoulder, scapula, vertebral line). Each of the 25 stimulus locations was presented in a pseudorandomized order while the participants closed their eyes. After the touch, the participant opened their eyes and marked the perceived location of the touch on the 3D rendering of their breast or back. The participant was encouraged to manipulate the 3D model freely to obtain the best view of their breast or back on each trial. The participant moved the cursor and clicked to indicate where they perceived the touch. Reported locations accumulated on the 3D model throughout the block (during which each location was touched once.) to encourage placement of markers in accurate relative positions but were removed at the start of each of three blocks (yielding a total of 3 repeats for each of the 25 locations).

Data Analysis

Added in the revised version: Spatial Precision

Removed in the revised version: Spatial Acuity

Earlier version: To quantify the spatial acuity at each test location, we fit a psychometric function (see EQ1) to the probability of judging the second touch as above the first versus the relative location of the two stimuli. From these functions, we estimated the distance from the reference point required to reliably discriminate the sensation (75% performance, see EQ2).

Revised version: To quantify the spatial precision at each test location, we fit a psychometric function (see EQ1) to the probability of judging the second touch as above the first versus the relative location of the two stimuli. From these functions, we estimated the distance from the reference point required to reliably discriminate the sensation (75% performance, see EQ2). (1) p( above )=11+e−k∗x

Earlier version: Where k is the growth term of the exponential function and x is the distance between the comparison point and the reference point.

Revised version: Where k is the growth term of the exponential function and x is the distance between the comparison point and the reference point. (2) jnd=log⁡(175−1)−k

Where k is the growth term computed in EQ1.

JNDs were then compared between regions using a 1-way ANOVA followed by a Tukey’s honest significant difference post-hoc test.

Added in the revised version: Quadrant localization

Added in the revised version: To determine if each participant individually performed the task at greater than chance levels as well as the group as a whole, a Monte-Carlo simulation was used. Across 10,000 simulations the appropriate number of trials were simulated in which each participant randomly guessed the location and the percent correct was computed across trials for each simulation. Significance was then computed within participant by comparing the observed performance against the simulated distribution and then applying Bonferroni correction. To then determine if the whole group outperformed chance, across each of the 10,000 simulations, the cross-participant average percent correct was computed and the observed percentage compared against the resultant value. The effect size was computed by comparing the observed average percent correct with the mean of the simulation divided by the standard deviation of the simulation.

Absolute localization task error and bias

To compute the error for the task, we considered the 3D position of both the stimuli and response. First, to compute the total error, we computed the 3D Euclidean distance between the stimulus location and the participants response on each trial and then averaged these values together. While the skin is not flat, over the relatively short distances between stimulus and response (typically <5 cm), Euclidean distance was sufficient as minimal curvature occurred over this range. To measure the bias for each point, the 3D position for all responses for a given point across blocks was averaged together and the Euclidean distance between the resultant point and the stimulus location was averaged (centroid). If responses were randomly distributed around the stimulus location, then the error of the average response would be minimal. If there was a consistent offset, then the bias error would be similar to the total error. Finally, to compute imprecision, we computed the mean Euclidean error between the centroid and individual responses.

Linear mixed models were used to assess significant effects using Matlab’s fitlme function. The formula used for each test can be found in Tables 1-4. While the fixed effects varied, participants were treated as random effects and the intercept and slope were allowed to vary for each.

Absolute localization task angular error

Earlier version: To compute the angle of the errors for each response we measured the four-quadrant inverse tangent (EQ3) between each point and response. Importantly, this was only considered in 2-dimensions as the variance of angles in the 3rd dimension (depth) were drastically dependent on breast size and the location of the point on the breast. Thus, the angle between the stimulus location and each response (as well as the centroid across responses) was computed. Then, for each stimulus the angle between the stimulus location and the reference point was computed (center angle) the difference in angle between said angle and the centroid angle was measured.

Revised version: To compute the angle of the errors for each response we measured the four-quadrant inverse tangent (EQ3) between each point and response. Importantly, this was only considered in 2-dimensions as the variance of angles in the 3rd dimension (depth) were drastically dependent on breast size and the location of the point on the breast. Thus, the angle between the stimulus location and each response (as well as the centroid across responses) was computed. Then, for each stimulus the angle between the stimulus location and the reference point was computed (center angle) the difference in angle between said angle and the centroid angle was measured. (3) θ=tan−1⁡(ΔyΔx)

Absolute localization task uniformity test

Earlier version: To test if the cross-participant biases were uniform or not we measured the range-normalized mean squared error (rnMSE, EQ4) between the average distribution of error angles across participants and the mean of the distribution, similar to measurements of residual variance. Then, to determine if the observed rnMSE was significant, we used a permutation test in which for each permutation (n = 1e5) we shuffled the angles for each participant before averaging across participants and computing the rnMSE to produce a null distribution of rnMSEs that might be expected if the participants responses were not systematic. Finally, we computed the proportion of simulations that had greater rnMSEs than the observed rnMSE to produce the p value which was considered significant if it was less than 0.05.

Revised version: To test if the cross-participant biases were uniform or not we measured the range-normalized mean squared error (rnMSE, EQ4) between the average distribution of error angles across participants and the mean of the distribution, similar to measurements of residual variance. Then, to determine if the observed rnMSE was significant, we used a permutation test in which for each permutation (n = 1e5) we shuffled the angles for each participant before averaging across participants and computing the rnMSE to produce a null distribution of rnMSEs that might be expected if the participants responses were not systematic.

Added in the revised version: Finally, we computed the proportion of simulations that had greater rnMSEs than the observed rnMSE to produce the p value which was considered significant if it was less than 0.05. (4) rnMSE=∑(x−x¯)22max{x}−min{x}

Absolute localization task vector strength

Earlier version: The vector strength for individual participants was computed using EQ5. Importantly, given the observed heteroskedasticity between distance and error, we did not incorporate the length of the vectors in the computation as this would have biased the computation towards points distal to the nipple for the breast but would not have influenced the back measurements. As the angular errors were distributed between −π and π, to test if the observed vector strength was significant, we used a Monte-Carlo simulation in which for each simulation (n = 1e5) we sampled angles from the range [-π, π] from a uniform distribution (n = 11) and computed the vector strength from each simulation. We then compared the observed vector strengths for each participant and location against the null distribution and considered it significant if the value exceeded the 95th percentile of the null distribution.

Revised version: The vector strength for individual participants was computed using EQ5. Importantly, given the observed heteroskedasticity between distance and error, we did not incorporate the length of the vectors in the computation as this would have biased the computation towards points distal to the nipple for the breast but would not have influenced the back measurements. As the angular errors were distributed between -π and π, to test if the observed vector strength was significant, we used a Monte-Carlo simulation in which for each simulation (n = 1e5) we sampled angles from the range [-π, π] from a uniform distribution (n = 11) and computed the vector strength from each simulation. We then compared the observed vector strengths for each participant and location against the null distribution and considered it significant if the value exceeded the 95th percentile of the null distribution. (5) vs=∑sin⁡(x)2+∑cos⁡(x)22n(x)